# BMI/CS 576 HW1
The objectives of this homework are to practice

* with the basic algorithms for sequence assembly
* reasoning about graphs and paths for the sequence assembly task

## HW policies
Before starting this homework, please read over the [homework policies](https://canvas.wisc.edu/courses/167969/pages/hw-policies) for this course.  In particular, note that homeworks are to be completed *individually*.

You are welcome to use any code from the weekly notebooks in your solutions to the HW.

## PROBLEM 1: The greedy algorithm for fragment assembly (60 points)
Write a function, `greedy_assemble`, that takes as input a list of read strings and uses the greedy fragment assembly algorithm to output a *single* superstring that contains all reads as substrings. You must use the graph-based (Hamiltonian path) version of the greedy algorithm. We will assume that:
1. we are assembling a single-stranded sequence and
2. that no read is a substring of any other read.

To keep things simple, for this homeowork, we will allow overlaps of any length (including zero).  In practice, for sequence assembly we would typically require some minimum overlap length.

### Tie-breaking criteria

For the purpose of making this algorithm deterministic, we must establish tiebreaking criteria for edges in the overlap graph that have the same weight. For two edges with the same weight, we will first choose the edge whose source vertex read is first in lexicographical order. If the source vertices are identical, then we choose the edge whose target vertex read is first in lexicographical order. For example, if e1 = ATCGGA → GGAT and e2 = ATCGGA → GGAA, we will attempt to use edge e2 first because GGAA < GGAT according to lexicographical order.

<!-- grade_id: cell-11ef4347de632758, points: 0, grade: false, solution: true -->
```python
# Code for PROBLEM 1
# You are welcome to develop your code as a separate Python module
# and import it here if that is more convenient for you.
def greedy_assemble(reads):
    """Assembles a set of reads using the graph-based greedy algorithm.
    
    Args:
        reads: a list of strings
    Returns:
        A string that is a superstring of the input reads
    """
    ###
    ### YOUR CODE HERE
    ###

```

Tests for `greedy_assemble` are provided at the bottom of this notebook.

## PROBLEM 2: Assembling a small subset of Ebola virus reads (10 points)

Included with this notebook is the file `ebola_reads.txt` which is small subset of the Illumina reads used to assemble the genome of an isolate of the Ebola virus, which caused a major epidemic in West Africa. 

Use your greedy assemble function to assemble these reads. Once correctly assembled, these reads form a short segment of the genome of this virus. To allow your assembler to succeed, the reads have been cleaned of errors and have have been oriented so that they all come from the same strand of the genome.  You may find the following function below of use, which produces a list of reads from the contents of a file.

```python
def read_strings_from_file(filename):
    return [line.rstrip() for line in open(filename)]
```

<!-- grade_id: cell-d08017f16c0169bf, points: 0, grade: false, solution: true -->
```python
###
### code for assembling the ebola reads
###

```

Once you have assembled the genomic segment, use the [BLAST](https://blast.ncbi.nlm.nih.gov/Blast.cgi) web service to search the NCBI database of proteins with your assembled sequence. You should use BLASTX with its default settings. Based on the results of your BLASTX search, which gene is contained within this genomic segment?

<!-- grade_id: problem2, points: 10, grade: true, solution: true -->
###
### *YOUR ANSWER TO PROBLEM 2 HERE*
###


### BONUS CHALLENGE
A subset of Illumina reads from an Ebola virus genome sequencing experiment that cover the entire genome are included in the file `ebola_full_genome_reads.txt`. Can you get your code to assemble these reads in under 2 minutes? Sorry, no extra credit here, just personal satisfaction!

## PROBLEM 3: SBH graphs and Eulerian paths (20 points) 
For the following strings, (i) give the k = 3 spectrum for the string, (ii) draw the SBH graph for the spectrum, (iii) give one Eulerian path and its corresponding string for the SBH graph, and (iv) show whether or not there exists an Eulerian path in the graph that corresponds to the original string.

(a) ATGGCCTGAATCC

(b) ATAGCCTAGCAAT

<!-- grade_id: problem3, points: 20, grade: true, solution: true -->
###
### *YOUR ANSWER TO PROBLEM 3 HERE*
###


## PROBLEM 4: Existence of DNA sequence with complete spectrum (10 points)
For any value of $k$, does there exist a DNA sequence that contains every possible $k$-mer exactly once?  Prove your answer.  *Hint: consider the SBH graph for the spectrum of such a sequence, should it exist*

<!-- grade_id: problem4, points: 10, grade: true, solution: true -->
###
### *YOUR ANSWER TO PROBLEM 4 HERE*
###


### Tests for problem 1

```python
def test_greedy_assemble_with_files(reads_filename, superstring_filename):
    reads = read_strings_from_file(reads_filename)
    [superstring] = read_strings_from_file(superstring_filename)
    assert greedy_assemble(reads) == superstring 
```

<!-- grade_id: greedy_assemble_returns_a_string, points: 20, grade: true, solution: false -->
```python
# TEST: greedy_assemble returns a string
sanity_test_reads = read_strings_from_file("tests/test_reads.txt")
assert isinstance(greedy_assemble(sanity_test_reads), str)
print("SUCCESS: greedy_assemble returns a string passed!")
```

<!-- grade_id: greedy_assemble_returns_a_superstring, points: 10, grade: true, solution: false -->
```python
# TEST: greedy_assemble returns a superstring
def is_superstring(s, reads):
    return all(read in s for read in reads)
assert is_superstring(greedy_assemble(sanity_test_reads), sanity_test_reads)
print("SUCCESS: greedy_assemble returns a superstring passed!")
```

<!-- grade_id: greedy_assemble_small_test_1, points: 5, grade: true, solution: false -->
```python
# TEST: greedy_assemble_small_test_1
small_test1_reads = ["GTT", "ATCTC", "CTCAA"]
assert greedy_assemble(small_test1_reads) == "ATCTCAAGTT"
print("SUCCESS: greedy_assemble_small_test_1 passed!")
```

<!-- grade_id: greedy_assemble_small_test_2, points: 5, grade: true, solution: false -->
```python
# TEST: greedy_assemble_small_test_2
small_test2_reads = ["CGAAG", "ATCGA", "AGAG", "GGG"]
assert greedy_assemble(small_test2_reads) == "ATCGAAGAGGG"
print("SUCCESS: greedy_assemble_small_test_2 passed!")
```

<!-- grade_id: greedy_assemble_small_test_3, points: 5, grade: true, solution: false -->
```python
# TEST: greedy_assemble_small_test_3
small_test3_reads = ["C", "T", "G", "A"]
assert greedy_assemble(small_test3_reads) == 'ACGT'
print("SUCCESS: greedy_assemble_small_test_3 passed!")
```

<!-- grade_id: greedy_assemble_large_test_1, points: 5, grade: true, solution: false -->
```python
# TEST: greedy_assemble large test 1
test_greedy_assemble_with_files("tests/large_test1_reads.txt", "tests/large_test1_superstring.txt")
print("SUCCESS: greedy_assemble large test 1 passed!")
```

<!-- grade_id: greedy_assemble_reads_7, points: 2, grade: true, solution: false -->
```python
# TEST: greedy_assemble_reads_7 (hidden)
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: greedy_assemble_reads_8, points: 2, grade: true, solution: false -->
```python
# TEST: greedy_assemble_reads_8 (hidden)
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: greedy_assemble_reads_9, points: 2, grade: true, solution: false -->
```python
# TEST: greedy_assemble_reads_9 (hidden)
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: greedy_assemble_reads_10, points: 2, grade: true, solution: false -->
```python
# TEST: greedy_assemble_reads_10 (hidden)
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: greedy_assemble_reads_11, points: 2, grade: true, solution: false -->
```python
# TEST: greedy_assemble_reads_11 (hidden)
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```