# Day 8 notebook

The objectives of this notebook are to practice with the concepts of

* dynamic programming
* memoization
* higher-order functions

## Knapsack problem

In today's activity, we will practice the concept of dynamic programming with a classical problem in Computer Sciences, the "Knapsack problem."  In the knapsack problem, we imagine that we are packing a knapsack (also known as a backpack) for a trip somewhere (e.g., a long backpacking trip in a remote area).  There are many items that we might like to pack in our knapsack, but it has a limited amount of space.  Suppose that there are $n$ items that we might want to pack, and the $i$th item has "size" $s_i$ and a "value" $v_i$.  The "size" is how much space the item will take up in the knapsack and the "value" is some measure of how important that item is for our trip (e.g., perhaps all of our items are food items and the value could be the number of calories in each food item).  We are given that the knapsack has a capacity $C$.

The Knapsack problem is thus to determine a subset of items to pack that 
1. can all fit in the knapsack together and 
2. maximizes the total value of the items in the knapsack.  
    
Mathematically, we wish to peform the following optimization, where $I$ is the set of $n$ items:

$\max_{I^* \subset I} \sum_{i \in I^*} v_i \\ \textrm{subject to the constraint}\ \sum_{i \in I^*} s_i \leq C$

### Solving Knapsack by dynamic programming
The Knapsack problem can be solved via a dynamic programming algorithm that breaks this problem down into the following subproblem. Let $F[capacity, num\_items]$ be the maximum value of items that could be packed into a knapsack with $capacity$ available space, only considering the first $num\_items$.  This subproblem can be solved recursively via the following recurrence:

$F[capacity, num\_items] = \left\{
\begin{array}{ll}
F[capacity, num\_items - 1], & \textrm{if } s_{num\_items} > capacity \\
\max\left\{
\begin{array}{l}
F[capacity, num\_items - 1], \\
F[capacity - s_{num\_items}, num\_items - 1] + v_{num\_items}
\end{array}\right.,& \textrm{if } s_{num\_items} \leq capacity
\end{array}\right.$

The intuition behind this recurrence is that we consider two cases for the item indexed by $num\_items$:
1.  we *exclude* this item from the knapsack (first term in the max) or
2.  we *include* this item in the knapsack, assuming it can fit in the given capacity (second term in the max)

Below we will practice with implementing a dynamic programming algorithm that uses this subproblem definition to solve the problem of finding the maximum total *value* of items that could be packed into the knapsack.  We will *not* address the related problem of determining the actual subset of items that achieves that maximum total value.

### PROBLEM 1: Recursive implementation (1 POINT)

We will start by simply implementing the solution via a recursive function, which implements the recurrence above.  This will not be an efficient algorithm, but it will be helpful to start with this as a baseline.  You are provided with some of the code already.

<!-- grade_id: cell-880617c3792b0553, points: 0, grade: false, solution: true -->
```python
def knapsack_score_recursive(capacity, sizes, values):
    """Solves the Knapsack problem via a recursive algorithm.

    Args:
        capacity: a number giving the amount of space in the knapsack.
        sizes: a list of the size of each possible item
        values: a list of the value of each possible item
    Returns:
        A number indicating the maximum total value of items that can be packed
        in the knapsack, with each item included in the knapsack at most once.
    """

    # A local helper function that computes the Knapsack subproblem (F[capacity, num_items])
    # described above in a recursive manner
    def subproblem(capacity, num_items):
        # base case
        if num_items == 0:
            return 0
        # case in which the item indexed by num_items (1-based) *cannot* fit in remaining capacity
        # only case #1 (exclusion) is possible
        elif sizes[num_items - 1] > capacity:
            return subproblem(capacity, num_items - 1)
        # case in which the item indexed by num_items (1-based) can fit in remaining capacity
        # either case #1 (exclusion) or case #2 (inclusion) is possible
        else:
            ###
            ### YOUR CODE HERE
            ###

    return subproblem(capacity, len(sizes))
```

<!-- grade_id: knapsack_score_recursive, points: 1, grade: true, solution: false -->
```python
# tests for knapsack_score_recursive
test_sizes = (5, 7, 3, 2)
test_values = (7, 8, 5, 3)
assert knapsack_score_recursive(0, test_sizes, test_values) == 0
assert knapsack_score_recursive(2, test_sizes, test_values) == 3
assert knapsack_score_recursive(3, test_sizes, test_values) == 5
assert knapsack_score_recursive(5, test_sizes, test_values) == 8
assert knapsack_score_recursive(12, test_sizes, test_values) == 16
assert knapsack_score_recursive(17, test_sizes, test_values) == 23
assert knapsack_score_recursive(17, [], []) == 0
assert knapsack_score_recursive(17, [18], [1]) == 0
print("SUCCESS: knapsack_score_recursive passed all tests!")
```

### PROBLEM 2: Memoized implementation (1 POINT)
To implement a dynamic programming approach to this problem, we will either need to use *memoization*, or compute values to the subproblems in a *bottom-up* approach.  Let's start with the memoization strategy, as it requires minimal modification to your recursive solution above.

Recall that the memoization strategy is to store the return values of a function in a lookup table (say a Python dictionary) indexed by the argument values for the function and to simply return the value in the lookup table when the function is called again with the same arguments.  We could do this by re-writing the `subproblem` code above, but it will be cleaner to do this in a general way via a *higher-order function* that can create a memoized version of *any* function. Below is the skeleton code for such a function that will return a memoized version of its input function, `func`, which we will assume takes exactly two arguments.  Complete the code below.

<!-- grade_id: cell-8b3fd59f0545dc5d, points: 0, grade: false, solution: true -->
```python
def memoize_2arg_func(func):
    """Creates a memoized version of a function that takes exactly two arguments.

    Args:
        func: a function that takes exactly two arguments
    Returns:
        A memoized version of the function func
    """
    # create a dictionary that will map arguments to return values of func
    outputs = dict()
    # define a new function that will be the memoized version
    def memoized(arg1, arg2):
        ###
        ### YOUR CODE HERE
        ###
    return memoized
```

<!-- grade_id: memoize_2arg_func, points: 1, grade: true, solution: false -->
```python
# tests for memoize_2arg_func
num_calls_to_test_func = 0 # keep track of how many times test_func is run
# a test function that returns the sum of two numbers
def test_func(x, y):
    global num_calls_to_test_func
    num_calls_to_test_func += 1
    return x + y
test_func = memoize_2arg_func(test_func)

assert test_func(1, 2) == 3
assert num_calls_to_test_func == 1
assert test_func(2, 3) == 5
assert num_calls_to_test_func == 2
assert test_func(1, 2) == 3 # repeated call
assert num_calls_to_test_func == 2
assert test_func(2, 3) == 5 # another repeated call
assert num_calls_to_test_func == 2

num_calls_to_test2_func = 0 # keep track of how many times test2_func is run
# a second test function that returns the length k prefix of string s
def test2_func(s, k):
    global num_calls_to_test2_func
    num_calls_to_test2_func += 1
    return s[:k]
test2_func = memoize_2arg_func(test2_func)

assert test2_func("madison", 3) == "mad"
assert num_calls_to_test2_func == 1
assert test2_func("madison", 3) == "mad" # repeated call
assert num_calls_to_test2_func == 1
assert test2_func("madison", 1) == "m"
assert num_calls_to_test2_func == 2
assert test2_func("madison", 1) == "m" # another repeated call
assert num_calls_to_test2_func == 2

print("SUCCESS: memoize_2arg_func passed all tests!")
```

Now paste in your `knapsack_score_recursive` code from above into the `knapsack_score_memoized` dynamic programming version of the algorithm below.  Note that `knapsack_score_memoized` creates a memoized version of the `subproblem` function using your `memoize_2arg_func` higher-order function. 

<!-- grade_id: cell-1f57595cad8de8b1, points: 0, grade: false, solution: true -->
```python
def knapsack_score_memoized(capacity, sizes, values):
    """Solves the Knapsack problem via a top-down dynamic programming algorithm.

    Args:
        capacity: a number giving the amount of space in the knapsack.
        sizes: a list of the size of each possible item
        values: a list of the value of each possible item
    Returns:
        A number indicating the maximum total value of items that can be packed
        in the knapsack, with each item included in the knapsack at most once.
    """
    
    def subproblem(capacity, num_items):
        if num_items == 0:
            return 0
        elif sizes[num_items - 1] > capacity:
            return subproblem(capacity, num_items - 1)
        else:
            ###
            ### YOUR CODE HERE
            ###

    subproblem = memoize_2arg_func(subproblem)
    return subproblem(capacity, len(sizes))
```

```python
# tests for knapsack_score_memoized (not graded)
test_sizes = (5, 7, 3, 2)
test_values = (7, 8, 5, 3)
assert knapsack_score_memoized(0, test_sizes, test_values) == 0
assert knapsack_score_memoized(2, test_sizes, test_values) == 3
assert knapsack_score_memoized(3, test_sizes, test_values) == 5
assert knapsack_score_memoized(5, test_sizes, test_values) == 8
assert knapsack_score_memoized(12, test_sizes, test_values) == 16
assert knapsack_score_memoized(17, test_sizes, test_values) == 23
assert knapsack_score_memoized(17, [], []) == 0
assert knapsack_score_memoized(17, [18], [1]) == 0
print("SUCCESS: knapsack_score_memoized passed all tests!")
```

#### Higher-order functions available in Python
Python is a rich functional language and offers a number of higher-order functions in its standard library.  For example, the [`map`](https://docs.python.org/3/library/functions.html#map) and [`filter`](https://docs.python.org/3/library/functions.html#filter) functions are commonly-used higher-order functions that you should consider using.  Other higher-order functions are found in the [`functools`](https://docs.python.org/3/library/functools.html) module. In particular, the [functools.lru_cache](https://docs.python.org/3/library/functools.html#functools.lru_cache) function is a more general implementation of the memoize function that you wrote above.

For higher-order functions that return functions, "[decorator](https://docs.python.org/3/glossary.html#term-decorator)" syntax can help.  For example,

```python
def test_func(x, y):
    return x + y
test_func = memoize_2arg_func(test_func)

# is more succinctly written with decorator syntax as
@memoize_2arg_func
def test_func(x, y):
    return x + y
```

### PROBLEM 3: Bottom-up implementation (1 POINT)
Lastly, we will implement the knapsack dynamic programming algorithm in a bottom-up style.  Here we first create a matrix that will store all of the results of the subproblems, and compute the subproblems iteratively, starting with the smallest subproblems first.  Below is some partially completed code that you are to fill in.

<!-- grade_id: cell-c4b37b13600a6a07, points: 0, grade: false, solution: true -->
```python
def knapsack_score_bottomup(capacity, sizes, values):
    """Solves the Knapsack problem via a bottom-up dynamic programming algorithm.

    Args:
        capacity: a number giving the amount of space in the knapsack.
        sizes: a list of the size of each possible item
        values: a list of the value of each possible item
    Returns:
        A number indicating the maximum total value of items that can be packed
        in the knapsack, with each item included in the knapsack at most once.
    """
    # inititialize a matrix (list of lists) to store the results
    # of the subproblem F[num_items, capacity]
    matrix = [[None] * (len(sizes) + 1) for i in range(capacity + 1)]
    # fill in first row
    for j in range(len(sizes) + 1):
        matrix[0][j] = 0
    # fill in first column
    for i in range(capacity + 1):
        matrix[i][0] = 0
    # fill in remaining elements
    for i in range(1, capacity + 1):
        for j in range(1, len(sizes) + 1):
            ###
            ### YOUR CODE HERE
            ###

    return matrix[capacity][len(sizes)]
```

<!-- grade_id: knapsack_score_bottomup, points: 1, grade: true, solution: false -->
```python
# tests for knapsack_score_bottomup
test_sizes = (5, 7, 3, 2)
test_values = (7, 8, 5, 3)
assert knapsack_score_bottomup(0, test_sizes, test_values) == 0
assert knapsack_score_bottomup(2, test_sizes, test_values) == 3
assert knapsack_score_bottomup(3, test_sizes, test_values) == 5
assert knapsack_score_bottomup(5, test_sizes, test_values) == 8
assert knapsack_score_bottomup(12, test_sizes, test_values) == 16
assert knapsack_score_bottomup(17, test_sizes, test_values) == 23
assert knapsack_score_bottomup(17, [], []) == 0
assert knapsack_score_bottomup(17, [18], [1]) == 0
print("SUCCESS: knapsack_score_bottomup passed all tests!")
```

### Comparing performance of the three implementations

With your three implementations of an algorithm to solve the knapsack problem, let's test their runtime to see how they perform relatively to each other.

```python
import timeit
capacity = 100
bigger_test_values = (6, 18, 28, 34, 4, 26, 38, 9, 36, 39, 33, 29, 3, 16, 37, 2, 24, 22, 30, 5)
bigger_test_sizes = (21, 12, 6, 0, 34, 1, 17, 22, 4, 32, 16, 33, 8, 20, 3, 18, 39, 36, 37, 19)
num_trials = 500 # number of times that timeit will run the function

print("Recursive:", timeit.timeit('knapsack_score_recursive(capacity, bigger_test_sizes, bigger_test_values)', 
                                  number=num_trials, globals=globals()), "seconds")
print(" Memoized:", timeit.timeit('knapsack_score_memoized(capacity, bigger_test_sizes, bigger_test_values)', 
                                  number=num_trials, globals=globals()), "seconds")
print("Bottom-up:", timeit.timeit('knapsack_score_bottomup(capacity, bigger_test_sizes, bigger_test_values)', 
                                  number=num_trials, globals=globals()), "seconds")

```

Do those timings match your expectations of how efficient these algorithms are relative to each other?  Write your thoughts in the cell below.