# Day 10 notebook

The objectives of this notebook are to practice

* representing characters in sequences as random variables
* computing marginal probabilities
* computing conditional probabilities
* checking for independence of random variables
* extracting features (e.g., gaps and mismatches) from an alignment
* comparing the distribution of some observations to a common distribution

## Modules used in this notebook

```python
from matplotlib import pyplot as plt    # for plotting
import collections                      # for helpful functions/classes, such as Counter and defaultdict
import fasta                            # for reading sequences in FASTA format
```

## Python functions to consider
In this activity, you will likely find some of the functions in the [collections](https://docs.python.org/3/library/collections.html) module to be extremely helpful.  

### Counter
The [Counter](https://docs.python.org/3/library/collections.html#collections.Counter) class is very helpful for counting the number of occurrences of things.  For example,

```python
print(collections.Counter("CATGATCTCATCGTACGCAACG"))
print(collections.Counter([("apple", 1), ("pear", 2), ("apple", 1), ("orange", 3)]))
```

### defaultdict
The [defaultdict](https://docs.python.org/3/library/collections.html#collections.defaultdict) class is very useful for automatically providing default values for keys in a dictionary.  For example,

```python
costs = collections.defaultdict(float)
# note that no errors are raised below
# if the key is not in the dictionary, it is added with a default value of 0.0
costs["apple"] += 1.3
costs["pear"] += 3
costs["apple"] += 0.5
print(costs)
```

## Data sets
Here are some of the data sets we will test and/or explore with this notebook.

### Weather/Flight data set
This is the dataset provided as an example in the lecture.

We will represent a probability distribution with a dictionary for which the keys are tuples (or other "sequence"-like objects, such as strings) representing a specific configuration of random variables and the values are the floating point values corresponding to the probabilities of the configurations.  Here is the joint probability distribution for the example given in the lecture in which there are two random variables: the current weather and the arrival status of a flight.

```python
weather_flight_joint_dist = {("sun",  "on-time"): 0.20,
                             ("rain", "on-time"): 0.20,
                             ("snow", "on-time"): 0.05,
                             ("sun",  "late"):    0.10,
                             ("rain", "late"):    0.30,
                             ("snow", "late"):    0.15}
```

We will represent a set of observations of a given system as a list of tuples (or other "sequence"-like objects, such as strings) representing a specific observation of all of the random variables in the system.  Such a set of observations can be thought of as a series of random draws from the joint probability distribution.  A set of observations for the weather/flight system is stored in the file `weather_flight_dataset.txt`, with each line corresponding to an observation.  We read this dataset from the file using the code below:

```python
weather_flight_dataset = [tuple(line.split()) for line in open("weather_flight_dataset.txt")]
```

Here are the first 6 observations of this dataset:

```python
print(*weather_flight_dataset[:6], sep="\n")
```

### Single base data set
This is a very simple dataset in which the observations are single DNA bases (sequences of length = 1).

```python
single_base_joint_dist = {'A': 0.3, 'C': 0.3, 'G': 0.2, 'T': 0.2}
single_base_dataset = ['A', 'C', 'A', 'T', 'G', 'T', 'C', 'A', 'G','C']
```

## Fruit fly splice site donor sequences

One of the two main dataset you will work with in this activity is a set of sequences from the fruit fly (*Drosophila melanogaster*) genome.  These sequences are taken from intervals of the genome that are called *donor splice sites*.  Donor splice sites are signals within genes that guide the processing of RNAs transcribed from those genes.  Specifically, donor splice site signals mark the 3' boundaries of *exons*, which are the intervals of genes that are retained in mature messenger RNAs.

The set of donor sequences are found in the file [fly_donor_sequences.txt](fly_donor_sequences.txt) within this workspace.  We will first read these sequences into a list:

```python
donors = [line.rstrip() for line in open("fly_donor_sequences.txt")]
```

Here are the first 6 donor sequences just to give you a sense of what they look like:

```python
print(*donors[:6], sep="\n")
```

## PROBLEM 1: Estimating joint distributions (1 POINT)
We will first write a function that takes a list of observations, where each observation is a sequence (e.g., a string or tuple) of random variable values and estimates a joint distribution for those random variables, by simply computing the frequency of each joint configuration of the random variables.  Consider using the [Counter](https://docs.python.org/3/library/collections.html#collections.Counter) class for this problem.

<!-- grade_id: cell-b6ffa320bbe479da, points: 0, grade: false, solution: true -->
```python
def estimate_joint_dist(observations):
    """Estimates the joint distribution from a list of observations.
    
    Args:
        observations: a list of observations.  Each observation should be a 
                      sequence-like object (e.g., a string or tuple), 
                      where the ith element of each sequence represents the 
                      observed value of the ith random variable.
    Returns:
        A dictionary with random variable configurations as keys and probabilities as values."""
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: estimate_joint_dist, points: 1, grade: true, solution: false -->
```python
# tests for estimate_joint_dist

# Because of floating point issues, we will use the function below to 
# round probabilities such that we can check for correctness to a certain 
# number of digits.
def round_dist(dist, digits=5):
    """Returns a new distribution with probabilities rounded to the specified number of digits."""
    return {key: round(value, digits) for key, value in dist.items()}

assert round_dist(estimate_joint_dist(single_base_dataset)) == single_base_joint_dist
assert round_dist(estimate_joint_dist(weather_flight_dataset)) == weather_flight_joint_dist
assert estimate_joint_dist(["apple"]) == {"apple": 1.0}
assert estimate_joint_dist([(22, "apple", 3.0), (3, "banana", 1.5)]) == {(22, "apple", 3.0): 0.5,
                                                                         (3, "banana", 1.5): 0.5}
print("SUCCESS: estimate_joint_dist passed all tests!")
```

### Examining the joint distribution of the donor sequences
Now use your function to estimate the joint distribution of the characters in the donor sequences:

```python
donor_joint_dist = estimate_joint_dist(donors)
```

Examine this distribution.  What is the most probabable donor sequence and what is its probability?

<!-- grade_id: cell-d596264fb3a93d63, points: 0, grade: false, solution: true -->
```python
###
### Your analysis here
###

```

## PROBLEM 2: Computing marginal distributions (1 POINT)
Now let us write a function that takes as input a joint distribution and computes the marginal distribution of one of the random variables.  Consider using the [defaultdict](https://docs.python.org/3/library/collections.html#collections.defaultdict) class for this problem.

<!-- grade_id: cell-e473d4613b5006c1, points: 0, grade: false, solution: true -->
```python
def compute_marginal_dist(joint_distribution, index):
    """Computes the marginal distribution of a single random variable given a joint distribution.
    
    Args:
        joint_distribution: A dictionary representing a joint probability distribution.  In the
                            dictionary, variable configurations are keys and probabilities are values
        index: an integer giving the index of the random variable for which the marginal 
               distribution is to be computed.  The index is with respect to the sequence objects
               that represent each variable configuration.
    Returns:
        A dictionary with the possible values of the specified random variable
        as keys and probabilities as values."""
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: compute_marginal_dist, points: 1, grade: true, solution: false -->
```python
# tests for compute_marginal_dist
assert round_dist(compute_marginal_dist(single_base_joint_dist, 0)) == single_base_joint_dist
assert round_dist(compute_marginal_dist(weather_flight_joint_dist, 0)) == {'rain': 0.5, 'snow': 0.2, 'sun': 0.3}
assert round_dist(compute_marginal_dist(weather_flight_joint_dist, 1)) == {'late': 0.55, 'on-time': 0.45}
assert round_dist(compute_marginal_dist(donor_joint_dist, 0)) == {'A': 0.36, 'C': 0.275, 'G': 0.196, 'T': 0.169}
print("SUCCESS: compute_marginal_dist passed all tests!")
```

### Examining the marginal distributions of the positions within donor sequences
We will now use your `compute_marginal_dist` function to compute the marginal distribution of each position within the donor sequences.

```python
donor_marginal_dists = [compute_marginal_dist(donor_joint_dist, i) for i in range(len(donors[0]))]
```

It will be useful to plot these marginal distributions in a series of bar plots.

```python
DNA_BASES = ("A", "C", "G", "T")

plt.figure(figsize=(14, 3))
for index, dist in enumerate(donor_marginal_dists):
    if index == 0:
        ax1 = plt.subplot(1, len(donor_marginal_dists), index + 1, title=index)
        plt.ylabel("probability")
    else:
        ax = plt.subplot(1, len(donor_marginal_dists), index + 1, sharey=ax1, title=index)
        plt.setp(ax.get_yticklabels(), visible=False)
    plt.bar(DNA_BASES, [dist.get(base, 0) for base in DNA_BASES])
plt.suptitle('Marginal distributions of bases at each donor position')
plt.subplots_adjust(wspace=0.15, top=0.8)
plt.show()
```

What strikes you about these disributions?

###
### *Your observations here*
###


## PROBLEM 3: Computing conditional distributions (1 POINT)
Now let us write a function that takes as input a joint distribution and computes the conditional distribution of one of the random variables given another.  The conditional distribution should be represented as a dictionary with the keys being values of the random variable being conditioned on, and the values being distributions over the values of the other random variable.

<!-- grade_id: compute_conditional_dist_implementation, points: 0, grade: false, solution: true -->
```python
def compute_conditional_dist(joint_distribution, index, condition_index):
    """Computes the conditional distribution of one random variable given another and a joint distribution.
    
    Args:
        joint_distribution: A dictionary representing a joint probability distribution.  In the
                            dictionary, variable configurations are keys and probabilities are values
        index: an integer giving the index of the random variable for which the conditional 
               distribution is to be computed. For example, if we are computing P(X|Y), 
               this is the index of X. The index is with respect to the sequence objects
               that represent each variable configuration. 
        condition_index: an integer giving the index of the random variable on which to condition.
               For example, if we are computing P(X|Y), this is the index of Y.
    Returns:
        A dictionary with the possible values of the conditioning random variable as keys
        and distributions (represented as dictionaries) over the targeted variable as values."""
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: compute_conditional_dist, points: 1, grade: true, solution: false -->
```python
# tests for compute_conditional_dist

def round_cond_dist(cond_dist, digits=5):
    """Returns a new conditional distribution with probabilities rounded to the specified number of digits."""
    return {key: round_dist(dist, digits) for key, dist in cond_dist.items()}

assert round_cond_dist(compute_conditional_dist(weather_flight_joint_dist, 1, 0)) == \
{'rain': {'late': 0.6,     'on-time': 0.4},
 'snow': {'late': 0.75,    'on-time': 0.25},
  'sun': {'late': 0.33333, 'on-time': 0.66667}}

print("SUCCESS: compute_conditional_dist passed all tests!")
```

### Examining conditional distributions of the positions within donor sequences
We will now use your `compute_conditional_dist` function to compute the conditional distribution of base at position 5 in the donor sequences given the base at position 7.

```python
donor_5_given_7_conditional_dist = compute_conditional_dist(donor_joint_dist, 5, 7)
```

Here is a series of plots showing these conditional distributions.

```python
plt.figure(figsize=(14, 3))
for index, (base, dist) in enumerate(sorted(donor_5_given_7_conditional_dist.items())):
    if index == 0:
        ax1 = plt.subplot(1, len(donor_5_given_7_conditional_dist), index + 1, title="position 7 = " + base)
        plt.ylabel("probability")
    else:
        ax = plt.subplot(1, len(donor_5_given_7_conditional_dist), index + 1, sharey=ax1, title="position 7 = " + base)
        plt.setp(ax.get_yticklabels(), visible=False)
    plt.bar(dist.keys(), dist.values())
plt.suptitle('Conditional distribution of the base at position 5 given the base at position 7')
plt.subplots_adjust(wspace=0.15, top=0.8)
plt.show()
```

Compare the above plots to the marginal distribution of the base at position 5 in the previous section.  Does it look like the base at position 5 is *independent* of the base at position 7?

###
### *Your thoughts here*
###


## Human-Chimp genomic alignment from chromosome 21
For the remainder of this activity we will analyze the distributions of various features within a pairwise sequence alignment.  In particular, we will consider the pairwise alignment of a pair orthologous segments of the human and chimpanzee genomes.  This alignment was produced using the [MUMmer](http://mummer.sourceforge.net/) genomic alignment software.  The human segment comes from bases 26727427-26958359 of chromosome 21, and is about 231,000 bases long.  We will first read this alignment from a file `human_chimp_chr21_block.fasta`, which is in FASTA format.

```python
aligned_seqs = fasta.read_sequences_from_fasta_file("human_chimp_chr21_block.fasta")
human_chimp_alignment = [aligned_seq for (seq_name, aligned_seq) in aligned_seqs]
```

The `slice_alignment` function below will be useful for looking at pieces of this alignment, which is very long.

```python
def slice_alignment(alignment, start_column, length):
    """Returns a slice of the alignment starting at start_column and with the given length."""
    return [aligned_seq[start_column: start_column + length] for aligned_seq in alignment]
```

And here is a view of a small 70-column slice of this alignment:

```python
slice_alignment(human_chimp_alignment, 107950, 70)
```

### Gap lengths

We will first examine the gaps in this alignment, and, in particular, the lengths of the gaps.

#### PROBLEM 4: Compute gap lengths in an alignment (1 POINT)

Write a function that takes as input a *single* sequence from an alignment that potentially has gaps (consecutive runs of space (`-`) characters), and returns a list of the lengths of the gaps it contains, in left to right order.

<!-- grade_id: cell-d84e54c47ae76294, points: 0, grade: false, solution: true -->
```python
def gap_lengths(s):
    """Computes a list of the lengths of the gaps in an aligned sequence
    
    Args:
        s: a string representing a single sequence from an alignment, which possibly
           contains gaps (consecutive runs of space characters, '-')
    Returns:
        A list of the lengths of the gaps in s (in left to right order)
    """
    ###
    ### YOUR CODE HERE
    ###


```

<!-- grade_id: gap_lengths, points: 1, grade: true, solution: false -->
```python
# tests for gap_lengths
assert gap_lengths("CTAGC---CGAT-C") == [3, 1]
assert gap_lengths("--CTAGC---CGAT-C") == [2, 3, 1]
assert gap_lengths("CTAGC---CGAT-C--") == [3, 1, 2]
assert gap_lengths("----------") == [10]
assert gap_lengths("CTAGCCGATC") == []
assert gap_lengths("") == []
print("SUCCESS: gap_lengths passed all tests!")
```

#### PROBLEM 5: Examination of the gap length distribution in a  human-chimp alignment (1 POINT)

We will now use your `gap_lengths` function to compile a list of all gaps from both sequences of the human-chimp alignment.

```python
human_chimp_gap_lengths = gap_lengths(human_chimp_alignment[0]) + gap_lengths(human_chimp_alignment[1])
print(human_chimp_gap_lengths)
```

Use `matplotlib.pyplot` to make a plot of the distribution of the gap lengths.  I recommend using the [`pyplot.hist`](https://matplotlib.org/api/_as_gen/matplotlib.pyplot.hist.html) on the list of lengths for this purpose.

<!-- grade_id: cell-3f69b4a0e9ed827a, points: 0, grade: false, solution: true -->
```python
###
### Code for plotting a histogram of the human_chimp_gap_lengths
###

```

Which common distribution described in the online lecture does this gap length distribution resemble?  Record your answer by assigning one of the following strings to the `gap_length_distribution` variable.

```python
common_distributions = {"geometric", "binomial", "multinomial", "gaussian", "exponential"}
```

<!-- grade_id: cell-454b7b6f77a393c9, points: 0, grade: false, solution: true -->
```python
###
### gap_length_distribution = ?
###

```

<!-- grade_id: gap_length_distribution, points: 1, grade: true, solution: false -->
```python
# tests for gap_length_count_distribution
assert gap_length_distribution in common_distributions, "The distribution must be one of the common distributions above"
###
### AUTOGRADER TEST - DO NOT REMOVE
###
print("SUCCESS: gap_length_distribution passed all visible tests!")
```

### Number of mismatches
Let's turn now to examining mismatches in this alignment.

#### PROBLEM 6: Compute the number of mismatches in an alignment (1 POINT)
Write a function that takes as input an alignment (represented as a list with two string elements) and returns the number of columns that have a mismatched pair of bases (ignoring gapped columns).

<!-- grade_id: cell-d90d4571bd527965, points: 0, grade: false, solution: true -->
```python
def num_mismatches(alignment):
    """Computes the number of mismatch columns (ignoring gapped columns) in alignment.
    
    Args:
        alignment: a pairwise alignment represented as a list with two string elements,
                   each string of the same length (the number of columns)
    Returns:
        The integer number of mismatch columns.
    """
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: num_mismatches, points: 1, grade: true, solution: false -->
```python
# tests for num_mismatches
test_align1 = ["CTAGGT",
               "C-ATCT"]
assert num_mismatches(test_align1) == 2
test_align2 = ["CTAG-T",
               "C-ATCT"]
assert num_mismatches(test_align2) == 1
test_align3 = ["TTTTTT",
               "AAAAAA"]
assert num_mismatches(test_align3) == 6
test_align4 = ["",
               ""]
assert num_mismatches(test_align4) == 0
print("SUCCESS: num_mismatches passed all tests!")
```

#### PROBLEM 7: Examination of the distribution of the number of mismatches in slices of a human-chimp alignment (1 POINT)

Now let's use your `num_mismatches` function to count the number of mismatches in slices of the human-chimp alignment.  We will slice the alignment into disjoint slices of 1000 columns each and count the number of mismatches in each slice.

```python
slice_width = 1000
mismatch_counts = [num_mismatches(slice_alignment(human_chimp_alignment, start, slice_width))
                   for start in range(0, len(human_chimp_alignment[0]), slice_width)]
```

Use `matplotlib.pyplot` to make a plot of the distribution of the mismatch counts.  I recommend using the [`pyplot.hist`](https://matplotlib.org/api/_as_gen/matplotlib.pyplot.hist.html) for this purpose.

<!-- grade_id: cell-29a7aadd9d08f257, points: 0, grade: false, solution: true -->
```python
###
### Code for plotting a histogram of the mismatch_counts
###

```

Which common distribution described in the online lecture does this mismatch count distribution resemble?  Record your answer by assigning one of the following strings to the `mismatch_count_distribution` variable.

<!-- grade_id: cell-69c3807049ef91b1, points: 0, grade: false, solution: true -->
```python
###
### mismatch_count_distribution = ?
###

```

<!-- grade_id: mismatch_count_distribution, points: 1, grade: true, solution: false -->
```python
# tests for mismatch_count_distribution
assert mismatch_count_distribution in common_distributions, "The distribution must be one of the common distributions above"
###
### AUTOGRADER TEST - DO NOT REMOVE
###
print("SUCCESS: mismatch_count_distribution passed all visible tests!")
```